C-C chemokine receptor 4 deficiency exacerbates early atherosclerosis in mice

Author response:

The following is the authors’ response to the previous reviews

Response to the reviewer #2 (Public review):

We greatly appreciate the reviewer’s high evaluation of our paper and helpful comments and suggestions.

Regarding in vivo Treg homing assay, we did not exclude doublets and dead cells from the analysis of Kaede-expressing Tregs migrated to the aorta, which may affect the results. We described this issue as the limitation of this study in the revised manuscript. Nonetheless, we believe the reliability of our findings because we repeated this experiment three times and obtained similar results.

There is no evidence to support the clinical relevance of our findings. Future clinical research on this topic is highly desired.

Response to the reviewer #3 (Public review):

We greatly appreciate the reviewer’s high evaluation of our paper and helpful comments and suggestions.

Despite the controversial role of Th17 cells in atherosclerosis, we understand the possible involvement of Th17 cells and the Th1 cell/Th17 cell balance in lymphoid tissues and aortic lesions in accelerated inflammation and atherosclerosis in Ccr4-/-Apoe-/- mice. Although we could not completely evaluate the changes in these immune responses in detail, future study may elucidate interesting mechanisms mediated by Th17 cell responses.

As the reviewer suggested, we understand that it is necessary to provide in vivo evidence for the Treg suppressive effects on DC activation. Based on the results of in vitro experiments, we described the discussion on the in vivo evidence in the revised manuscript.

We understand methodological limitations for flow cytometric analysis of immune cells in the aorta and in vivo Treg homing assay. We described this issue as the limitation of this study in the revised manuscript. Regarding in vivo Treg homing assay, we statistically re-analyzed the combined data from multiple experiments and observed a tendency toward reduction in the proportion of CCR4-deficient Kaede-expressing Tregs in the aorta of recipient Apoe-/- mice, though there was no statistically significant difference in the migratory capacity of CCR4-intact or CCR4-deficient Kaede-expressing Tregs. Accordingly, we toned down our claim that CCR4 expression on Tregs plays a critical role in mediating Treg migration to the atherosclerotic aorta under hypercholesterolemia.

The reviewer requested us to evaluate aortic inflammation in Ccr4-/-Apoe-/- mice injected with CCR4-intact or CCR4-deficient Tregs. However, we think that this experiment will provide marginal information because Treg transfer experiments in Apoe-/- mice have already shown the protective role of CCR4 in Tregs against aortic inflammation and early atherosclerosis.

Recommendations for the authors:

Reviewer #2 (Recommendations for the authors):

(1) #1 and #2: CD103 and CD86 expression should be discussed on the text and not only in the response to reviewer.

In accordance with the reviewer’s suggestion, we added a discussion on the downregulated CD103 expression in peripheral LN Tregs and upregulated CD86 expression on DCs in Ccr4-/-Apoe-/- mice in the discussion section in the revised manuscript.

(2) #5: Authors response is not satisfactory. No gate percentage is shown. As it currently is, the difference in the number of cells shown in the figure could be due to differences in events recorded. Furthermore, the gate strategy is not thorough. Considering the very low frequency of Kaede + cells detected, it is crucial to properly exclude doublets and dead cells.

Authors reported a dramatic difference in Kaede + Tregs cells in the aorta across experiments. This could be addressed by normalization followed by appropriate statistical analysis (One sample t-test).

The data shown is not strong enough to conclude that there is a reduced migration to the aorta.

We understand the importance of reviewer’s suggestion. We described the percentage of Kaede+ Tregs in the aorta of Apoe-/- mice receiving transfer of Kaede-expressing CCR4-intact or CCR4-deficient Tregs in Figure 5I.

As the reviewer pointed out, we understand that it would be important to properly exclude doublets and dead cells in in vivo Treg homing assay. However, it is difficult for us to resolve this issue because we need to perform the same experiments again which will require a great number of additional mice and substantial amount of time. We deeply regret that these important experimental procedures were not performed. We described this issue as the limitation of this study.

In accordance with the reviewer’s suggestion, we re-analyzed the combined data from multiple experiments using one-sample t-test. We observed a tendency toward reduction in the proportion of CCR4-deficient Kaede-expressing Tregs in the aorta of recipient Apoe-/- mice, though there was no statistically significant difference in the migratory capacity of CCR4-intact or CCR4-deficient Kaede-expressing Tregs. By modifying the corresponding descriptions in the manuscript, we toned down our claim that CCR4 expression on Tregs plays a critical role in mediating Treg migration to the atherosclerotic aorta under hypercholesterolemia.

(3) #8: There are still several not shown data

In accordance with the reviewer’s suggestion, we showed the data on the responses of Tregs and effector memory T cells in 8-week-old wild-type or Ccr4-/- mice and Ccr4 mRNA expression in Tregs and non-Tregs from Apoe-/- or Ccr4-/-Apoe-/- mice in Supplementary Figures 4 and 7.

Reviewer #3 (Recommendations for the authors):

(1) Issue 1. For future studies, I recommend not omitting viability controls during cell staining. Removal of dead cells and doublets should always be included during the gating strategy to avoid undesirable artefacts, especially when analysing less-represented cell populations. According to your previous report (ref #40), I agree that isotype controls were unnecessary using the same staining protocol. FMO controls should always be included in flow cytometry analysis (not mentioned in the methodology description and ref#40).

As the reviewer suggested, we understand that it would be important to properly exclude dead cells and doublets and to prepare FMO controls in flow cytometric analysis. We deeply regret that these important experimental procedures were not performed. We described this issue as the limitation of this study.

(2) Issue 3. Although Th17's role in atherosclerosis remains controversial, the data obtained in this work could provide valuable insights if discussed appropriately. As noted in my public review, I found it noteworthy that ROR γ t+ cells represented around 13% of effector TCD45+CD3+CD4+ lymphocytes in the aorta of Apoe-/- mice while Th1 less than 5% (Fig 4H and F, respectively). I recognise that differences in cell staining sensibility and robustness for different transcription factors may influence these percentages. However, analysing how CCR4 deficiency influences the Th1/TI h17 balance would yield interesting data, similar to what was done for the Th1/Treg ratio.

Considering the higher proportion of Th17 cells than Th1 or Th2 cells in atherosclerotic aorta, we understand the importance of reviewer’s suggestion. However, we could not evaluate the effect of CCR4 deficiency on the Th1/Th17 balance in aorta because we did not perform flow cytometric analysis of aortic Th1 and Th17 cells in the same mice. Meanwhile, we could examine the Th1/Th17 balance in peripheral lymphoid tissues by flow cytometry. We found a significant increase in the Th1/Th17 ratio in the peripheral LNs of Ccr4-/-Apoe-/- mice, while there were no changes in its ratio in the spleen or para-aortic LNs of these mice, which limits the contribution of the Th1/Th17 balance to exacerbated atherosclerosis. We showed these data below.

(3) Issue 4. I appreciate the authors for sharing data on the flow cytometry analysis of Tregs in para-aortic LNs of Apoe-/- and Ccr4-/- Apoe-/- mice, which would have been included as a Supplementary figure. These results reinforce the notion that Treg dysfunction in CCR4-deficient mice may not be due to the downregulation of regulatory cell surface receptors.

We showed the data on the expression of CTLA-4, CD103, and PD1 in Tregs in the para-aortic LNs of Apoe-/- and Ccr4-/-Apoe-/- mice in Supplementary Figure 8.

(4) Issue 5. I agree that CD4+ T cell responses are substantially regulated by DCs. While CD80 and CD86 on DC primarily serve as costimulatory signals for T-cell activation, cytokines secreted by DCs are primordial signals for determining the differentiation phenotype of effector Th cells. Since the analysis of DC phenotype in lymphoid tissues of Apoe-/- and Ccr4-/- Apoe-/- mice could not be addressed in this study, it is not possible to differentiate which processes may be mainly affected by CCR4-deficiency during CD4+ T cell activation. In this scenario, and considering in vitro studies, the results suggest a possible role of CCR4 in controlling the extent of activation of CD4+T cells rather than shifting the CD4+T cell differentiation profile in peripheral lymphoid tissues, where a predominant Th1 profile was already established in Apoe-/- mice. Therefore, I advise caution when concluding about shifts in CD4+ T cell responses.

We thank the reviewer for providing us thoughtful comments. As the reviewer pointed out, we understand that we should carefully interpret the mechanisms for the shift of CD4+ T cell responses by CCR4 deficiency.

(5) Regarding migration studies in the revised manuscript. I fully understand that Treg transference assays are challenging. The results do not suggest that CCR4 was critical for Treg migration to lymphoid tissues in the conditions assayed. Concerning migration to the aorta, I found the results inconclusive since the authors mention that: i) there was a dramatic difference in the absolute numbers of Kaede-expressing Tregs that migrated to the aorta impairing statistical analysis; ii) the number of Kaede-expressing Tregs that migrated to the aorta was extremely low; iii) dead cells and doublets were not removed in the flow cytometry analysis. In this context, I do not agree with the following statements and recommend revising them:

- "CCR4 deficiency in Tregs impaired their migration to the atherosclerotic aorta" (lines 36-7),

- "…we found a significant reduction in the proportion of CCR4 deficient Kaede-expressing Tregs in the aorta of recipient Apoe-/- mice" (lines 356-7),

- "CCR4 expression on Tregs regulates the development of early atherosclerosis by....... mediating Treg migration to the atherosclerotic aorta" (lines 409-411),

- "…we found that CCR4 expression on Tregs is critical for regulating atherosclerosis by mediating their migration to the atherosclerotic aorta" (lines 437-438),

- "CCR4 protects against early atherosclerosis by mediating Treg migration to the aorta.... (lines 464-465),

- "We showed that CCR4 expression on Tregs is critical for ...... mediating Treg migration to the atherosclerotic aorta" (503-505).

We understand the importance of the reviewer’s suggestion. We described this issue as the limitation of this study. In accordance with the reviewer’s suggestion, we modified the above descriptions and toned down our claim that CCR4 expression on Tregs plays a critical role in mediating Treg migration to the atherosclerotic aorta under hypercholesterolemia.

(6) Line 206: Mention the increased expression of CD86 by DCs

We mentioned this result in the revised manuscript. We also added a discussion on the upregulated CD86 expression on DCs in Ccr4-/-Apoe-/- mice in the discussion section in the revised manuscript.

(7) Lines 304-305. According to Fig 4F-H, a selective accumulation of Th1 cells seems to have occurred only in the aorta, coinciding with a higher Th1/Treg ratio. No selective accumulation of Th1 cells was observed in para-aortic lymph nodes. These results could be clarified.

We modified the above description in the revised manuscript.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: Chronic inflammation via dysregulation of T cell immune responses is critically involved in the pathogenesis of atherosclerotic cardiovascular disease. Improving the balance between proinflammatory T cells and anti-inflammatory regulatory T cells (Tregs) may be an attractive approach for treating atherosclerosis. Although C-C chemokine receptor 4 (CCR4) has been shown to mediate the recruitment of T cells to inflamed tissues, its role in atherosclerosis is unclear. Here, we show that genetic deletion of CCR4 in hypercholesterolemic mice accelerates the development of early atherosclerotic lesions characterized by an inflammatory plaque phenotype. This was associated with the augmentation of proinflammatory T helper type 1 (Th1) cell responses in peripheral lymphoid tissues, para-aortic lymph nodes, and atherosclerotic aorta. Mechanistically, CCR4 deficiency in Tregs impaired their suppressive function and migration to the atherosclerotic aorta and augmented Th1 cell-mediated immune responses through defective regulation of dendritic cell function, which accelerated aortic inflammation and atherosclerotic lesion development. Thus, we revealed a previously unrecognized role for CCR4 in controlling the early stage of atherosclerosis via Treg-dependent regulation of proinflammatory T cell responses. Our data suggest that CCR4 is an important negative regulator of atherosclerosis.

Revised version: Chronic inflammation via dysregulation of T cell immune responses is critically involved in the pathogenesis of atherosclerotic cardiovascular disease. Improving the balance between proinflammatory T cells and anti-inflammatory regulatory T cells (Tregs) may be an attractive approach for treating atherosclerosis. Although C-C chemokine receptor 4 (CCR4) has been shown to mediate the recruitment of T cells to inflamed tissues, its role in atherosclerosis is unclear. Here, we show that genetic deletion of CCR4 in hypercholesterolemic mice accelerates the development of early atherosclerotic lesions characterized by an inflammatory plaque phenotype. This was associated with the augmentation of proinflammatory T helper type 1 (Th1) cell responses in peripheral lymphoid tissues, para-aortic lymph nodes, and atherosclerotic aorta. Mechanistically, CCR4 deficiency in Tregs impaired their suppressive function and tended to inhibit their migration to the atherosclerotic aorta, and subsequently augmented Th1 cell-mediated immune responses through defective regulation of dendritic cell function, which accelerated aortic inflammation and atherosclerotic lesion development. Thus, we revealed a previously unrecognized role for CCR4 in controlling the early stage of atherosclerosis via Treg-dependent regulation of proinflammatory T cell responses. Our data suggest that CCR4 is an important negative regulator of atherosclerosis.

Introduction

Severe cardiovascular diseases, including ischemic heart disease and stroke, occur in patients with atherosclerosis and are major causes of mortality worldwide. Despite state-of-the-art intensive treatment, patients at high risk of atherosclerotic diseases have considerable residual risk, which could be partly related to vascular inflammatory responses.1 Notably, recent clinical trials have provided evidence that anti-inflammatory therapies could be potentially novel approaches for preventing cardiovascular diseases in patients with the past disease history, although overall mortality has not improved.2,3

Chronic aortic inflammation via T cell-mediated immune dysregulation has been shown to be critically involved in the development of atherosclerosis.4 Single-cell proteomic and transcriptomic approaches have revealed that activated and differentiated T cells are the major immune cells infiltrating human carotid artery plaques,5 indicating the involvement of T cell-mediated immunoinflammatory responses in atherosclerotic plaque development in humans. A recent clinical trial provided evidence for the increased risk of atherosclerotic cardiovascular disease events in cancer patients treated with immune checkpoint inhibitors,6 confirming the critical role of effector T cell (Teff) immune responses in the development of human atherosclerotic disease. Teffs including T helper type 1 (Th1) cells play a proatherogenic role by producing proinflammatory cytokines including interferon (IFN)-γ.7,8 A recent translational study demonstrated that anti-CXC-motif-chemokine receptor 3 autoantibodies, which reflect Th1 cell-mediated responses, can be a novel biomarker and an inflammatory risk factor for cardiovascular morbidity and mortality beyond traditional risk factors, indicating the possible proatherogenic role of Th1 cells in humans.9 On the other hand, several subsets of regulatory T cells (Tregs) including forkhead box P3 (Foxp3)-expressing Tregs play an anti-atherogenic role by suppressing Teff-mediated immunoinflammatory responses or reducing plasma atherogenic lipoprotein levels.10–12 A number of approaches to dampening proatherogenic Th1 cell responses or augmenting atheroprotective Treg responses have been reported to be effective for the treatment and prevention of atherosclerosis in experimental mouse models.13 Despite accumulating evidence for the critical role of the Th1 cell/Treg balance in atherosclerosis, how the balance of these T cell populations is regulated in lymphoid tissues and atherosclerotic lesions to control atherosclerosis is unclear.

Various immune cells including T cells and monocytes migrate to the aorta via interactions between chemokines and their specific chemokine receptors and are critically involved in the process of atherosclerosis.14 However, few reports have identified a chemokine system that prevents atherosclerosis by improving the Th1 cell/Treg balance in lymphoid tissues and atherosclerotic lesions. C-C chemokine receptor 4 (CCR4) is expressed on several T cell subsets, including T helper type 2 (Th2) cells, T helper type 17 (Th17) cells, skin-homing T cells, and Tregs, but not on proatherogenic Th1 cells.15 CCR4 is a highly specific receptor for two CC chemokine ligands CCL17 (thymus- and activation-regulated chemokine) and CCL22 (macrophage-derived chemokine)16 which were shown to be expressed in mesenteric lymph nodes (LNs) in a mouse model of inflammatory bowel disease17 and lung tissues affected by allergic airway inflammation.18 These chemokines play a role in guiding CCR4+ T cells to the inflammatory sites with abundant expression of these chemokines.19,20 Previous experimental studies have suggested that the CCL17/CCL22– CCR4 axes protect against inflammatory autoimmune diseases partly by promoting Treg accumulation in target tissues.17,18 The expression of CCL17 and CCL22 was also observed in human and mouse atherosclerotic lesions.21–23 Although the disruption of the CCL17/CCL22–CCR4 axes did not affect the development of advanced atherosclerotic lesions or the proportion of Tregs in peripheral lymphoid tissues in atherosclerosis-prone mice fed a high-cholesterol diet,22,24 its effect on the early stage of atherosclerosis and Treg responses has not been investigated.

Here, we investigated the role of CCR4 in the development of early atherosclerotic lesions and the underlying mechanisms in CCR4-deficient (Ccr4-/-) mice on an apolipoprotein E-deficient (Apoe-/-) background fed a standard chow diet, with a particular focus on CD4+ T cell immune responses.

Results

CCR4 is predominantly expressed on CD4+Foxp3+ Tregs, and the ligands CCL17 and CCL22 are expressed in peripheral LNs and atherosclerotic lesions

To evaluate the expression levels of CCR4 on CD4+ T cells under normocholesterolemic or hypercholesterolemic conditions, we performed flow cytometric analysis of peripheral LNs, spleen, and para-aortic LNs from wild-type, Apoe-/-, and Ccr4-/-Apoe-/- mice. Notably, CCR4 expression was expressed on approximately 15-25% of CD4+Foxp3+ Tregs from wild-type or Apoe-/- mice (Supplementary Figure 1A through 1C), while CD4+Foxp3- non-Tregs from these mice expressed CCR4 at markedly lower levels, suggesting that CCR4 is predominantly expressed on CD4+Foxp3+ Tregs and that hypercholesterolemia does not affect CCR4 expression. We also examined CCR4 expression on CD4+Foxp3+ Tregs in the atherosclerotic aorta of Apoe-/- mice, and consistently found that CCR4 expression in CD4+Foxp3+ Tregs was much higher than that in CD4+Foxp3- non-Tregs (Supplementary Figure 1D).

As CCL17 and CCL22, known as specific ligands for CCR4, are highly expressed by dendritic cells (DCs) in LNs,25,26 we examined their expression in the peripheral LNs of wild-type, Apoe-/-, and Ccr4-/-Apoe-/- mice by immunohistochemistry. As expected, clear expression of these chemokines was observed in the peripheral LNs of wild-type, Apoe-/-, and Ccr4-/-Apoe-/- mice (Supplementary Figure 2). We also examined the expression of CCL17 and CCL22 in the aorta of these mice by immunohistochemistry. CCL17 expression was detected in the atherosclerotic lesions of Apoe-/- and Ccr4-/-Apoe-/- mice, and some lesional MOMA-2+ macrophages expressed this chemokine (Supplementary Figure 3A). However, CCL17 expression was not detected in the aorta of wild-type mice without atherosclerotic plaques (Supplementary Figure 3A). Another ligand CCL22 was modestly expressed in atherosclerotic lesions and partially colocalized with lesional MOMA-2+ macrophages in Apoe-/- or Ccr4-/-Apoe-/- mice, while its expression was not detected in the aorta of wild-type mice (Supplementary Figure 3B).

Together, these findings suggest that Tregs may migrate to peripheral LNs and atherosclerotic lesions partly via the CCL17/CCL22–CCR4 axes under hypercholesterolemia.

CCR4 deficiency accelerates the development of early atherosclerotic lesions characterized by an inflammatory plaque phenotype

To investigate the effect of CCR4 deficiency on the development of early atherosclerosis, we analyzed the atherosclerotic lesions of 18-week-old Apoe-/- and Ccr4-/-Apoe-/- mice fed a standard chow diet. Ccr4-/-Apoe-/- mice developed normally without any spontaneous inflammatory disease. Notably, compared with Apoe-/- mice, Ccr4-/-Apoe-/- mice exhibited a significant increase in atherosclerotic lesion size in the aortic sinus (aortic sinus mean plaque area: control Apoe-/- mice: 1.46 ± 0.50 ×105 μm2 versus Ccr4-/-Apoe-/-: 2.04 ± 0.82 ×105 μm2, P<0.05; Figure 1A). In parallel with the cross-sectional studies, we performed en face analysis of thoracoabdominal aortas, revealing no significant difference in aortic plaque burden between the 2 groups (Figure 1B). There were no significant differences in body weight or plasma lipid profile between the 2 groups (Supplementary Table 1).

C-C chemokine receptor 4 (CCR4) deficiency accelerates the development of early atherosclerotic lesions characterized by an inflammatory plaque phenotype.

A, Representative photomicrographs of Oil Red O staining and quantitative analysis of atherosclerotic lesion area at 5 different levels and the average area in the aortic sinus of 18-week-old apolipoprotein E-deficient (Apoe-/-) mice (n=24) or CCR4-deficient mice on an Apoe-/- background (Ccr4-/-Apoe-/- ; n=23). B, Representative photomicrographs of Oil Red O staining and quantitative analysis of atherosclerotic lesion area in the aorta of 18-week-old Apoe-/- (n=15) or Ccr4-/-Apoe-/- mice (n=15). C-E, Representative sections and quantitative analyses of MOMA-2+ macrophages (C), CD4+ T cells (D), and collagen (E) in the aortic sinus. Arrowheads indicate the CD4+ T cells. n=10 per group. F, mRNA expression of pro- or anti-inflammatory cytokines and helper T cell-associated transcription factors in aorta. The expression levels of the target genes were normalized so that the mean values in Apoe-/- mice were set to 1. n=8 to 10 per group. Eighteen-week-old Apoe-/- or Ccr4-/-Apoe-/- mice were used for all experiments. Black bars represent 50, 200, or 500 μm as described. Data points represent individual animals. Horizontal bars represent means. Error bars indicate s.d. *P<0.05, **P<0.01; Mann-Whitney U-test: A and F Il1b; 2-tailed Student’s t-test: C, D, E, and F Il6, Tbx21, and Rorc.

To determine the effect of CCR4 deficiency on plaque components, we performed immunohistochemical studies of atherosclerotic lesions in the aortic sinus. Notably, compared with those of Apoe-/- mice, the atherosclerotic lesions of Ccr4-/-Apoe-/- mice showed a 20% increase in macrophage accumulation (Figure 1C) and a marked 42% increase in CD4+ T cell infiltration (Figure 1D). In addition, we performed immunohistochemical analysis of Foxp3+ Tregs in atherosclerotic lesions using an anti-Foxp3 antibody. However, few Foxp3+ Tregs were found within the plaques of Apoe-/- or Ccr4-/-Apoe-/- mice (data not shown). The proportion of collagen in the aortic sinus plaques of Ccr4-/-Apoe-/- mice was significantly lower than that in the aortic sinus plaques of Apoe-/- mice (Figure 1E). These findings on atherosclerotic plaque size and components collectively suggest that CCR4 plays a role in preventing the development of early atherosclerotic lesions and inducing a less inflammatory plaque phenotype. To further evaluate aortic immunoinflammatory responses, we analyzed the mRNA expression of pro- and anti-inflammatory cytokines and transcription factors specific for Tregs or helper T cell subsets in atherosclerotic aorta by quantitative reverse transcription PCR. The mRNA expression of proinflammatory cytokines (Il1b and Il6), Th1-related Tbx21, and Th17-related Rorc was markedly upregulated in the aorta of Ccr4-/-Apoe-/- mice, indicating augmented proatherogenic immune responses in the atherosclerotic aorta (Figure 1F). The mRNA expression of the Treg-specific transcription factor Foxp3 was undetectable.

Collectively, these data suggest that CCR4 deficiency promotes the accumulation of inflammatory cells and proinflammatory immune responses in the aorta, leading to augmented development of early atherosclerotic lesions in the aortic root.

CCR4 deficiency augments Teff immune responses in peripheral lymphoid tissues

We examined the mechanisms by which CCR4 deficiency accelerates early atherosclerosis by focusing on changes in systemic T cell responses, including those involving CD4+Foxp3+ Tregs and CD4+Foxp3- non-Tregs, in peripheral lymphoid tissues. The frequencies of CD4+Foxp3+ Tregs and CD4+CD44highCD62Llow effector memory T cells were significantly higher in the spleen of 8- or 18-week-old Ccr4-/-Apoe-/- mice than in those of age-matched Apoe-/- mice (Figure 2A and 2B). Similar tendency was observed for the peripheral LNs of these mice, although there was no difference in the frequency of CD4+CD44highCD62Llow effector memory T cells between 8-week-old Apoe-/- and Ccr4-/-Apoe-/- mice (Figure 2A and 2B). The absolute numbers of CD4+Foxp3+ Tregs and CD4+CD44highCD62Llow effector memory T cells were significantly higher in the peripheral LNs of 8- or 18-week-old Ccr4-/-Apoe-/- mice than those of age-matched Apoe-/- mice (Figure 2A and 2B). The absolute number of CD4+CD44highCD62Llow effector memory T cells was also higher in the spleen of 18-week-old Ccr4-/-Apoe-/- mice than that of age-matched Apoe-/- mice (Figure 2B).

C-C chemokine receptor 4 (CCR4) deficiency augments effector T cell immune responses in peripheral lymphoid tissues.

Earlier version: A and B, Representative flow cytometric analysis of CD4+ forkhead box P3 (Foxp3)+ regulatory T cells (Tregs) (A) and CD4+CD44highCD62Llow effector memory T cells (B) in the spleen of 8-week-old apolipoprotein E-deficient (Apoe−/−) mice or CCR4-deficient mice on an Apoe−/− background (Ccr4−/−Apoe− The graphs represent the total numbers and proportions of CD4+Foxp3+ Tregs (A) and CD4+CD44highCD62Llow effector memory T cells (B) in the peripheral lymph nodes (LNs) and spleen of 8-or 18-week-old Apoe−/− or Ccr4−/−Apoe−/− mice. n=9 to 10 per group. C and D, The graphs represent proportions of Ki-67 positive cells among CD4+Foxp3+ Tregs (C) and CD4+Foxp3− non-Tregs (D) in the peripheral LNs and spleen of 8-week-old Apoe−/− or Ccr4−/− Apoe−/− mice, as assessed by flow cytometry. n=13 per group. E, Expression levels of activation-associated molecules cytotoxic T lymphocyte-antigen-4 (CTLA-4) and CD103 were analyzed by gating on CD4+Foxp3+ Tregs in the peripheral LNs of 8-or 18-week-old Apoe−/− or Ccr4−/− Apoe−/− mice. n=9 to 10 per group. F, mRNA expression of Treg-associated markers in splenic Tregs from 8-week-old Apoe−/− or Ccr4−/−Apoe−/− mice. n=8 G, Expression levels of activation-associated molecules CTLA-4 and CD103 were analyzed by gating on CD4+Foxp3− non-Tregs in the LNs of 8-or 18-week-old Apoe−/− or Ccr4−/−Apoe−/− mice. n=9 to 10 per group. H, mRNA expression of activation-associated molecules in n-Tregs from 8-week-old Apoe−/− or Ccr4−/−Apoe−/− mice. n=8 per group. The expression levels of the target genes were normalized so that the es in Apoe−/− mice were set to 1 (F, H). Data points represent individual animals. Horizontal bars represent means. Error bars indicate s.d. *P<0.05, **P<0.01; Mann-Whitney U-test: A second (8w) from the left, B second (8w) and third (8w) from the left, C left, and H Cd44 and Cd103; 2-tailed Student’s t-test: A first, second (18w), and third from the left, B first, second (18w), third (18w), and fourth from the left, C right, D, E, G, and H indicates mean fluorescence intensity.

Revised version: A and B, Representative flow cytometric analysis of CD4+ forkhead box P3 (Foxp3)+ regulatory T cells (Tregs) (A) and CD4+CD44highCD62Llow effector memory T cells (B) in the spleen of 8-week-old apolipoprotein E-deficient (Apoe-/-) mice or CCR4-deficient mice on an Apoe-/- background (Ccr4-/-Apoe-/-). The graphs represent the total numbers and proportions of CD4+Foxp3+ Tregs (A) and CD4+CD44highCD62Llow effector memory T cells (B) in the peripheral lymph nodes (LNs) and spleen of 8- or 18- week-old Apoe-/- or Ccr4-/-Apoe-/- mice. n=9 to 10 per group. C and D, The graphs represent the proportions of Ki-67 positive cells among CD4+Foxp3+ Tregs (C) and CD4+Foxp3- non-Tregs (D) in the peripheral LNs and spleen of 8-week-old Apoe-/- or Ccr4-/-Apoe-/- mice, as assessed by flow cytometry. n=13 per group. E, Expression levels of activation-associated molecules cytotoxic T lymphocyte-associated antigen-4 (CTLA-4) and CD103 were analyzed by gating on CD4+Foxp3+ Tregs in the peripheral LNs of 8- or 18-week-old Apoe-/- or Ccr4-/- Apoe-/- mice. n=9 to 10 per group. F, mRNA expression of Treg-associated markers in splenic Tregs from 8-week-old Apoe-/- or Ccr4-/-Apoe-/- mice. n=8 per group. G, Expression levels of activation-associated molecules CTLA-4 and CD103 were analyzed by gating on CD4+Foxp3- non-Tregs in the peripheral LNs of 8- or 18-week-old Apoe-/- or Ccr4-/-Apoe-/- mice. n=9 to 10 per group. H, mRNA expression of activation-associated molecules in splenic non-Tregs from 8-week-old Apoe-/- or Ccr4-/-Apoe-/- mice. n=8 per group. The expression levels of the target genes were normalized so that the mean values in Apoe-/- mice were set to 1 (F, H). Data points represent individual animals. Horizontal bars represent means. Error bars indicate s.d. *P<0.05, **P<0.01; Mann-Whitney U-test: A second (8w) from the left, B second (8w) and third (8w) from the left, C left, and H Cd44 and Cd103; 2-tailed Student’s t-test: A first, second (18w), and third from the left, B first, second (18w), third (18w), and fourth from the left, C right, D, E, G, and H Ctla4. MFI indicates mean fluorescence intensity.

Earlier version: Similar results were obtained for normocholesterolemic wild-type mice (data not shown), suggesting that the expansion of Tregs and effector memory T cells in Ccr4−/−Apoe−/− mice is independent of hypercholesterolemia. We evaluated the proliferative capacity of CD4+Foxp3+ Tregs and CD4+Foxp3− non-Tregs by analyzing Ki-67 expression using flow cytometry and found that the proportions of Ki-67-positive Tregs and non-Tregs were markedly higher in the peripheral LNs and spleen of Ccr4−/−Apoe−/− mice than in those of Apoe−/− mice (Figure 2C and 2D). To determine the effect of CCR4 deficiency on T cell development in the thymus, we performed flow cytometric analysis of thymocytes from 4-week-old Apoe−/− or Ccr4−/−Apoe−/− mice and found no difference in the development of thymic T cells between the 2 groups (Supplementary Figure 4A). In line with a previous report in normocholesterolemic mice,27 there was no difference in the frequency of thymic CD4+Foxp3+ Tregs between the 2 groups (Supplementary Figure 4B). These data indicate that the increased numbers of CD4+Foxp3+ Tregs and CD4+Foxp3− non-Tregs in peripheral LNs are due to their enhanced proliferative capacity but not to their promoted development in the thymus. We also analyzed other immune cells in the spleen of Apoe−/− and Ccr4−/−Apoe−/− mice by flow cytometry and found no major differences in the proportions or activation-associated molecule expression between the 2 groups (Supplementary Figure 5).

Revised version: Similar results were obtained for normocholesterolemic wild-type mice (Supplementary Figure 4A and 4B), suggesting that the expansion of Tregs and effector memory T cells in Ccr4-/-Apoe-/- mice is independent of hypercholesterolemia. We evaluated the proliferative capacity of CD4+Foxp3+ Tregs and CD4+Foxp3- non-Tregs by analyzing Ki-67 expression using flow cytometry and found that the proportions of Ki-67-positive Tregs and non-Tregs were markedly higher in the peripheral LNs and spleen of Ccr4-/-Apoe-/- mice than in those of Apoe-/- mice (Figure 2C and 2D). To determine the effect of CCR4 deficiency on T cell development in the thymus, we performed flow cytometric analysis of thymocytes from 4-week-old Apoe-/- or Ccr4-/-Apoe-/- mice and found no difference in the development of thymic T cells between the 2 groups (Supplementary Figure 5A). In line with a previous report in normocholesterolemic mice,27 there was no difference in the frequency of thymic CD4+Foxp3+ Tregs between the 2 groups (Supplementary Figure 5B). These data indicate that the increased numbers of CD4+Foxp3+ Tregs and CD4+Foxp3- non-Tregs in peripheral LNs are due to their enhanced proliferative capacity but not to their promoted development in the thymus. We also analyzed other immune cells in the spleen of Apoe-/- and Ccr4-/-Apoe-/- mice by flow cytometry and found no major differences in the proportions or activation-associated molecule expression between the 2 groups, except for CD86 expression on DCs which was upregulated in Ccr4-/-Apoe-/- mice (Supplementary Figure 6).

Earlier version: To determine the activation and function of CCR4-deficient CD4+Foxp3+ Tregs, we investigated the expression of their activation– and function-associated molecules in the peripheral LNs of 8-or 18-week-old Apoe−/− or Ccr4−/−Apoe−/− mice by flow cytometry. Notably, CCR4 deficiency had no major effect on the expression of cytotoxic T lymphocyte-associated antigen-4 (CTLA-4) or CD103 in CD4+Foxp3+ Tregs (Figure 2E). We also analyzed the mRNA expression of Ccr4, activation-or function-associated molecules (Foxp3, Ctla4, Cd103, Tnfrsf18, Il10, and Tgfb), and major chemokine receptors (Ccr5, Ccr6, Ccr7, and Ccr8) in splenic Tregs by quantitative reverse transcription PCR. The Ccr4 mRNA expression in Tregs from Ccr4−/−Apoe−/− mice was less than the detectable levels (data not shown). In line with the data on peripheral LNs, there were no significant differences in the mRNA expression of these molecules between the 2 groups (Figure 2F; Supplementary Figure 6A).

Revised version: To determine the activation and function of CCR4-deficient CD4+Foxp3+ Tregs, we investigated the expression of their activation- and function-associated molecules in the peripheral LNs of 8- or 18-week-old Apoe-/- or Ccr4-/-Apoe-/- mice by flow cytometry. Notably, CCR4 deficiency had no major effect on the expression of cytotoxic T lymphocyte-associated antigen-4 (CTLA-4) or CD103 in CD4+Foxp3+ Tregs (Figure 2E). We also analyzed the mRNA expression of Ccr4, activation- or function-associated molecules (Foxp3, Ctla4, Cd103, Tnfrsf18, Il10, and Tgfb), and major chemokine receptors (Ccr5, Ccr6, Ccr7, and Ccr8) in splenic Tregs by quantitative reverse transcription PCR. The Ccr4 mRNA expression in Tregs from Ccr4-/-Apoe-/- mice was less than the detectable levels (Supplementary Figure 7A). In line with the data on peripheral LNs, there were no significant differences in the mRNA expression of these molecules between the 2 groups (Figure 2F; Supplementary Figure 7A).

Earlier version: In contrast to the results on CD4+Foxp3+ Tregs, the expression of activation marker CTLA-4 in peripheral LN CD4+Foxp3− non-Tregs was higher in Ccr4−/−Apoe−/− mice than in Apoe−/− mice, although the expression of another activation marker CD103 was not altered by CCR4 deficiency (Figure 2G). We analyzed the mRNA expression of Ccr4, activation-associated molecules (Ctla4, Cd44, Cd69, and Cd103), and representative chemokine receptors (Ccr1, Ccr5, Ccr6, Ccr7, Ccr8, Cxcr3, and Cx3cr1) in splenic non-Tregs by quantitative reverse transcription PCR. The Ccr4 mRNA expression in non-Tregs from Ccr4−/−Apoe−/− mice was less than the detectable levels (data not shown). We found a marked increase in the mRNA expression of activation-associated molecules (Ctla4, Cd44, and Cd103) in splenic non-Tregs from Ccr4−/−Apoe−/− mice (Figure 2H). We also observed upregulated mRNA expression of various chemokine receptors including Th1-related Ccr1, Cxcr3, and Cx3cr1, Th2-related Ccr8, and Th17-related Ccr6 in splenic non-Tregs from Ccr4−/−Apoe−/− mice (Supplementary Figure 6B), indicating that CCR4 deficiency may promote the activation and expansion of helper T cells and facilitate their migration to atherosclerotic lesions. There were no significant differences in the mRNA expression of other chemokine receptors (Ccr5 and Ccr7) in splenic non-Tregs between the 2 groups (Supplementary Figure 6B).

Revised version: In contrast to the results on CD4+Foxp3+ Tregs, the expression of activation marker CTLA-4 in peripheral LN CD4+Foxp3- non-Tregs was higher in Ccr4-/-Apoe-/- mice than in Apoe-/- mice, although the expression of another activation marker CD103 was not altered by CCR4 deficiency (Figure 2G). We analyzed the mRNA expression of Ccr4, activation-associated molecules (Ctla4, Cd44, Cd69, and Cd103), and representative chemokine receptors (Ccr1, Ccr5, Ccr6, Ccr7, Ccr8, Cxcr3, and Cx3cr1) in splenic non-Tregs by quantitative reverse transcription PCR. The Ccr4 mRNA expression in non-Tregs from Ccr4-/-Apoe-/- mice was less than the detectable levels (Supplementary Figure 7B). We found a marked increase in the mRNA expression of activation-associated molecules (Ctla4, Cd44, and Cd103) in splenic non-Tregs from Ccr4-/-Apoe-/- mice (Figure 2H). We also observed upregulated mRNA expression of various chemokine receptors including Th1-related Ccr1, Cxcr3, and Cx3cr1, Th2-related Ccr8, and Th17-related Ccr6 in splenic non-Tregs from Ccr4-/-Apoe-/- mice (Supplementary Figure 7B), indicating that CCR4 deficiency may promote the activation and expansion of helper T cells and facilitate their migration to atherosclerotic lesions. There were no significant differences in the mRNA expression of other chemokine receptors (Ccr5 and Ccr7) in splenic non-Tregs between the 2 groups (Supplementary Figure 7B).

Considering the accelerated early atherosclerosis observed in Ccr4-/-Apoe-/- mice, these results indicate that augmented Teff immune responses may affect atherosclerosis more strongly than the increase in Tregs in peripheral lymphoid tissues.

CCR4 deficiency promotes proinflammatory CD4+ T cell immune responses in peripheral lymphoid tissues

To determine whether CCR4 deficiency affects CD4+ T cell immune responses and polarization, we examined cytokine secretion from CD4+ T cells by intracellular cytokine staining. The fraction of IFN-γ-producing Th1 cells in the peripheral LNs was significantly higher in Ccr4-/-Apoe-/- mice than in Apoe-/- mice, while there were no differences in the fractions of other CD4+ T cell subsets including interleukin (IL)-4-producing Th2 cells, IL-10-producing CD4+ T cells, and IL-17-producing Th17 cells between the 2 groups (Figure 3A). In line with this, the fraction of splenic Th1 cells was higher in Ccr4-/-Apoe-/- mice than in Apoe-/- mice, although the proportion of IL-17-producing Th17 cells was also higher in Ccr4-/-Apoe-/- mice than in Apoe-/- mice (Figure 3B).

C-C chemokine receptor 4 (CCR4) deficiency promotes proinflammatory CD4+ T cell immune responses in peripheral lymphoid tissues.

A and B, The graphs represent the frequencies of interferon (IFN)-γ+, interleukin (IL)-4+, IL-10+, and IL-17+ CD4+ T cells in the peripheral lymph nodes (LNs) (A) and spleen (B) of 8-week-old apolipoprotein E-deficient (Apoe-/-) mice or CCR4-deficient mice on an Apoe-/- background (Ccr4-/-Apoe-/-). n=10 to 11 per group. C and D, Purified splenic CD4+ T cells from 8-week-old Apoe-/- or Ccr4-/-Apoe-/- mice were stimulated with plate-bound anti-CD3 and soluble anti-CD28 antibodies in vitro. The levels of various cytokines and chemokines in pooled cell supernatants from 8 mice in each group were determined semiquantitatively by a cytokine array kit (C). Data are representative of two independent experiments. Cytokine concentrations in the cell supernatants were measured by ELISA (D). n=8 per group. Data points represent individual animals. Horizontal bars represent means. Error bars indicate s.d. *P<0.05, **P<0.01; Mann-Whitney U-test: A, B first from the left, and D second from the left; 2-tailed Student’s t-test: B fourth from the left and D first, third, and fourth from the left.

We semiquantitatively analyzed the production of various cytokines or chemokines by splenic CD4+ T cells stimulated with plate-bound anti-CD3 and anti-CD28 antibodies using a cytokine array kit. Compared with those from Apoe-/- mice, splenic CD4+ T cells from Ccr4-/-Apoe-/- mice secreted more Th1-related cytokine IFN-γ, Th2-related cytokine IL-13, Th17-related cytokine IL-17, and various inflammation-related cytokines and chemokines (Figure 3C). ELISA analysis confirmed that the IFN-γ and IL-17 levels in the cell supernatants of Ccr4-/-Apoe-/- mice were much higher than those in the cell supernatants of Apoe-/- mice (Figure 3D). Although the cytokine levels of Th2-related cytokine IL-4 and Treg-related anti-inflammatory cytokine IL-10 were below the detectable levels by cytokine array analysis, ELISA analysis revealed higher IL-4 and IL-10 production in splenic CD4+ T cells from Ccr4-/-Apoe-/- mice than in those from Apoe-/- mice (Figure 3D), which may be compensatory immune responses to the proinflammatory T cell responses caused by CCR4 deficiency.

CCR4 deficiency promotes Th1 cell responses in para-aortic LNs and atherosclerotic aorta

Next, we investigated the mechanisms by which CCR4 deficiency accelerates early atherosclerosis by focusing on local immune responses in para-aortic LNs and atherosclerotic aorta. Consistent with the peripheral LN data, the frequency and number of CD4+Foxp3+ Tregs were significantly higher in the para-aortic LNs of Ccr4-/-Apoe-/- mice than in those of Apoe-/- mice (Figure 4A). The number of CD4+CD44highCD62Llow effector memory T cells was significantly higher in the para-aortic LNs of Ccr4-/-Apoe-/- mice than in those of Apoe-/- mice, while their frequency did not differ between the 2 groups (Figure 4B).

C-C chemokine receptor 4 (CCR4) deficiency promotes T helper type 1 (Th1) cell responses in para-aortic lymph nodes (LNs) and atherosclerotic aorta.

Earlier version: A and B, Representative flow cytometric analysis of CD4+ forkhead box P3 (Foxp3)+ regulatory T cells (Tregs) (A) and CD4+CD44high CD62Llow effector memory T cells (B) in para-aortic LNs. The graphs represent the total numbers and proportions of CD4+Foxp3+ Tregs (A) and CD4+CD44highCD62Llow effector memory T cells (B) in para-aortic LNs. n=9 to 10 per group. C and D, mRNA expression of Treg-associated markers in Tregs (C) and mRNA expression of activation or helper T cell-associated molecules in non-Tregs (D) in para-aortic LNs. The expression levels of the target genes were normalized so that the mean values in apolipoprotein E-deficient (Apoe−/−) mice were set to 1. Tregs or non-Tregs purified from pooled para-aortic LNs of 9 to 10 mice were analyzed as a sample. n=4 per group. E, The graphs represent the frequencies of interferon (IFN)-γ+, interleukin (IL)-4+, IL-10+, and IL-17+ CD4+ T cells in para-aortic LNs. n=12 per group. F-H, Representative flow cytometric analysis of T-box expressed in T cells (T-bet) (F), GATA3 (G), and retinoic acid-related orphan receptor gamma t (RORγt) (H) expression in aortic CD3+CD4+CD45+ T cells. The graphs represent the frequencies of T-bet+ (F), GATA3+ (G), and RORγt+ (H) cells among aortic CD3+CD4+CD45+ T cells. n=9 per group. I, Representative flow cytometric analysis of Foxp3 expression in aortic CD3+CD4+CD45+ T cells. The graph represents the frequency of Foxp3+ Tregs among aortic CD3+CD4+CD45+ T cells. n=9 per group. J, The graph represents the ratio of CD4+T-bet+ Th1 cells to CD4+Foxp3+ Tregs (Th1 cell/Treg ratio). n=9 per group. Pooled aortic lymphoid cells from 2 mice were analyzed as a sample. Eighteen-week-old Apoe−/− or CCR4-deficient mice on an Apoe−/− background (Ccr4−/−Apoe−/−) were used for all experiments. Data points represent individual animals (A, B, and E) or individual pooled samples (C, D, F-J). Horizontal bars represent means. Error bars indicate s.d. *P<0.05, **P<0.01; Mann-Whitney U-test: D Cd103 and J; 2-tailed Student’s t-test: A, B, D Tbx21, E, and F. left.

Revised version: A and B, Representative flow cytometric analysis of CD4+ forkhead box P3 (Foxp3)+ regulatory T cells (Tregs) (A) and CD4+CD44highCD62Llow effector memory T cells (B) in para-aortic LNs. The graphs represent the total numbers and proportions of CD4+Foxp3+ Tregs (A) and CD4+CD44highCD62Llow effector memory T cells (B) in para-aortic LNs. n=9 to 10 per group. C and D, mRNA expression of Treg-associated markers in Tregs (C) and mRNA expression of activation or helper T cell-associated molecules in non-Tregs (D) in para-aortic LNs. The expression levels of the target genes were normalized so that the mean values in apolipoprotein E-deficient (Apoe-/-) mice were set to 1. Tregs or non-Tregs purified from pooled para-aortic LNs of 9 to 10 mice were analyzed as a sample. n=4 per group. E, The graphs represent the frequencies of interferon (IFN)-γ+, interleukin (IL)-4+, IL-10+, and IL-17+ CD4+ T cells in para-aortic LNs. n=12 per group. F-H, Representative flow cytometric analysis of T-box expressed in T cells (T-bet) (F), GATA3 (G), and retinoic acid-related orphan receptor gamma t (RORγt) (H) expression in aortic CD3+CD4+CD45+ T cells. The graphs represent the frequencies of T-bet+ (F), GATA3+ (G), and RORγt+ (H) cells among aortic CD3+CD4+CD45+ T cells. n=9 per group. I, Representative flow cytometric analysis of Foxp3 expression in aortic CD3+CD4+CD45+ T cells. The graph represents the frequency of Foxp3+ Tregs among aortic CD3+CD4+CD45+ T cells. n=9 per group. J, The graph represents the ratio of CD4+T-bet+ Th1 cells to CD4+Foxp3+ Tregs (Th1 cell/Treg ratio). n=9 per group. Pooled aortic lymphoid cells from 2 mice were analyzed as a sample. Eighteen-week-old Apoe-/- or CCR4-deficient mice on an Apoe-/- background (Ccr4-/-Apoe-/-) were used for all experiments. Data points represent individual animals (A, B, and E) or individual pooled samples (C, D, F-J). Horizontal bars represent means. Error bars indicate s.d. *P<0.05, **P<0.01; Mann-Whitney U-test: D Cd103 and J; 2-tailed Student’s t-test: A, B, D Tbx21, E, and F.

Earlier version: In line with the data on peripheral lymphoid tissues, there were no significant differences in the expression of activation-or function-associated molecules or chemokine receptors in para-aortic LN Tregs between the 2 groups (Figure 4C; Supplementary Figure 7A). Notably, the expression of Cd103 and Th1-related Tbx21 was upregulated in para-aortic LN non-Tregs from Ccr4−/−Apoe−/− mice, while the expression of other activation-or helper T cell-associated molecules or chemokine receptors was unaltered (Figure 4D; Supplementary Figure 7B). Flow cytometric analysis of helper T cell subsets in para-aortic LNs revealed that the fractions of IFN-γ-producing Th1 cells, IL-4-producing Th2 cells, and IL-17-producing Th17 cells were significantly higher in Ccr4−/−Apoe−/− mice than in Apoe−/− mice, suggesting augmented proinflammatory CD4+ T cell immune responses by CCR4 deficiency (Figure 4E).

Revised version: In line with the data on peripheral lymphoid tissues, there were no significant differences in the expression of activation- or function-associated molecules or chemokine receptors in para-aortic LN Tregs between the 2 groups (Figure 4C; Supplementary Figures 8 and 9A). Notably, the expression of Cd103 and Th1-related Tbx21 was upregulated in para-aortic LN non-Tregs from Ccr4-/-Apoe-/- mice, while the expression of other activation- or helper T cell-associated molecules or chemokine receptors was unaltered (Figure 4D; Supplementary Figure 9B). Flow cytometric analysis of helper T cell subsets in para-aortic LNs revealed that the fractions of IFN-γ-producing Th1 cells, IL-4-producing Th2 cells, and IL-17-producing Th17 cells were significantly higher in Ccr4-/-Apoe-/- mice than in Apoe-/- mice, suggesting augmented proinflammatory CD4+ T cell immune responses by CCR4 deficiency (Figure 4E).

We explored the infiltration of T-box expressed in T cells (T-bet)-expressing Th1 cells, GATA3-expressing Th2 cells, and retinoic acid-related orphan receptor gamma t (RORγt)-expressing Th17 cells in atherosclerotic aorta by flow cytometry. Strikingly, the frequency of aortic T-bet-expressing Th1 cells was markedly higher in Ccr4-/-Apoe-/- mice than in Apoe-/- mice, while the proportions of Th2 and Th17 subsets did not differ between the 2 groups (Figure 4F through 4H). We next analyzed the accumulation of CD4+Foxp3+ Tregs in the atherosclerotic aorta by flow cytometry and found no difference in the proportion of CD4+Foxp3+ Tregs between the 2 groups (Figure 4I). Importantly, we found a marked increase in the ratio of T-bet-expressing Th1 cells to CD4+Foxp3+ Tregs (Th1 cell/Treg ratio) in the atherosclerotic aorta of Ccr4-/-Apoe-/- mice (Figure 4J), indicating Th1-skewed immune responses in the atherosclerotic lesions of Ccr4-/-Apoe-/- mice.

Earlier version: Collectively, these results suggest that CCR4 deficiency promotes the accumulation of proatherogenic Th1 cells in para-aortic LNs and atherosclerotic aorta and shifts the Th1 cell/Treg balance toward Th1 cell responses, leading to exacerbated aortic inflammation and early atherosclerosis.

Revised version: Collectively, these results suggest that CCR4 deficiency promotes the selective accumulation of proatherogenic Th1 cells in atherosclerotic aorta and the accumulation of various helper T cell subsets including Th1 cells in para-aortic LNs and shifts the Th1 cell/Treg balance toward Th1 cell responses in atherosclerotic aorta, leading to exacerbated aortic inflammation and early atherosclerosis.

Earlier version: CCR4 expression on Tregs regulates Th1 cell responses and mediates Treg migration to the atherosclerotic aorta

Revised version: CCR4 expression on Tregs regulates Th1 cell responses and may mediate Treg migration to the atherosclerotic aorta

Given that CCR4 plays a crucial role in attenuating immunoinflammatory responses in autoimmune or allergic disease via the modulation of Treg function,17,18 we speculated that despite the expansion and unaltered expression of activation- or function-associated molecules in Tregs, these cells might be dysfunctional. To determine whether CCR4 deficiency affects the suppressive function of Tregs in hypercholesterolemia, we performed an in vitro suppression assay. Interestingly, CCR4 deficiency significantly impaired the suppressive function of Tregs isolated from hypercholesterolemic mice (Figure 5A), indicating that CCR4 expression on Tregs may be important for the regulation of proinflammatory immune responses and the development of early atherosclerosis.

C-C chemokine receptor 4 (CCR4) expression on regulatory T cells (Tregs) regulates T helper type 1 cell responses and mediates Treg migration to the aorta.

Earlier version: A, The suppressive function of Tregs was assessed by evaluating the proliferation of carboxyfluorescein diacetate succinimidyl ester (CFSE)-labeled conventional T cells (Tconv) cocultured with Tregs from apolipoprotein E-deficient (Apoe−/−) mice or CCR4-deficient mice on an Apoe−/− background (Ccr4−/−Apoe−/−). Data are presented as the results of triplicate wells and are representative of 2 independent experiments. Data are expressed as the mean±s.d. B and C, CD80 and CD86 expression in live splenic dendritic cells (DCs) after 2 days of coculture with Tconv from Apoe−/− mice, or a mixture of Tregs from Apoe−/− or Ccr4−/−Apoe−/− mice and Tconv from Apoe−/− mice in the presence of an anti-CD3 antibody. Data points represent the results of quintuplicate wells. Data are representative of 2 independent experiments. B and D, Tconv from Apoe−/− mice and DCs were cocultured with or without Tregs from Apoe−/− or Ccr4−/−Apoe−/− mice in the presence of an anti-CD3 antibody. Interferon (IFN)-γ concentrations in cell supernatants were measured by ELISA. Data points represent the results of sextuplicate wells. E, Eighteen-week-old Apoe−/− mice fed a high-cholesterol diet for 10 weeks received transfer of Tregs from Apoe−/−Kaede-Tg or Ccr4−/−Apoe−/−Kaede-Tg mice, and the accumulation of Kaede+ Tregs in the peripheral lymphoid tissues and aorta was analyzed by flow cytometry 20 hours later. F-I, Representative flow cytometric analysis and the proportions of Kaede+ Tregs among CD4+ T cells in the peripheral lymph nodes (LNs) (F), spleen (G), para-aortic LNs (H), and aorta (I) of Apoe−/− mice that received Apoe−/−Kaede+ Tregs or Ccr4−/−Apoe−/−Kaede+ Tregs. n=9 to 10 per group (F-H). Pooled aortic lymphoid cells from 5 mice in each group were used for analysis. Data are representative of 3 independent experiments with similar results (I). Data points represent individual animals (F-H) or individual pooled samples (I). Horizontal bars represent means. Error bars indicate s.d. *P<0.05, **P<0.01; 1-way ANOVA followed by Tukey’s multiple comparisons test: C and D; 2-way ANOVA followed by Tukey’s multiple comparisons test: A. MFI indicates mean fluorescence intensity.

Revised version: A, The suppressive function of Tregs was assessed by evaluating the proliferation of carboxyfluorescein diacetate succinimidyl ester (CFSE)-labeled conventional T cells (Tconv) cocultured with Tregs from apolipoprotein E-deficient (Apoe-/-) mice or CCR4-deficient mice on an Apoe-/- background (Ccr4-/-Apoe-/-). Data are presented as the results of triplicate wells and are representative of 2 independent experiments. Data are expressed as the mean±s.d. B and C, CD80 and CD86 expression in live splenic dendritic cells (DCs) after 2 days of coculture with Tconv from Apoe-/- mice, or a mixture of Tregs from Apoe-/- or Ccr4-/-Apoe-/- mice and Tconv from Apoe-/- mice in the presence of an anti-CD3 antibody. Data points represent the results of quintuplicate wells. Data are representative of 2 independent experiments. B and D, Tconv from Apoe-/- mice and DCs were cocultured with or without Tregs from Apoe-/- or Ccr4-/-Apoe-/- mice in the presence of an anti-CD3 antibody. Interferon (IFN)-γ concentrations in cell supernatants were measured by ELISA. Data points represent the results of sextuplicate wells. E, Eighteen-week-old Apoe-/- mice fed a high-cholesterol diet for 10 weeks received transfer of Tregs from Apoe-/-Kaede-Tg or Ccr4-/-Apoe-/-Kaede-Tg mice, and the accumulation of Kaede+ Tregs in the peripheral lymphoid tissues and aorta was analyzed by flow cytometry 20 hours later. F-I, Representative flow cytometric analysis and the proportions of Kaede+ Tregs among CD4+ T cells in the peripheral lymph nodes (LNs) (F), spleen (G), para-aortic LNs (H), and aorta (I) of Apoe-/- mice that received Apoe-/-Kaede+ Tregs or Ccr4-/-Apoe-/-Kaede+ Tregs. n=9 to 10 per group (F-H). Pooled aortic lymphoid cells from 5 mice in each group were used for analysis. The results are presented as the mean ±s.d. of 3 independent experiments (I). Data points represent individual animals (F-H) or individual pooled samples (I). Horizontal bars represent means. Error bars indicate s.d. *P<0.05, **P<0.01; 1-way ANOVA followed by Tukey’s multiple comparisons test: C and D; 2-way ANOVA followed by Tukey’s multiple comparisons test: A. P=0.09; one sample t-test: I. MFI indicates mean fluorescence intensity.

Interaction with DCs is well-known as one of the core suppressive mechanisms by which Tregs control excessive immune responses. A previous report demonstrated that the cell–cell contacts between Tregs and CCL22-deficient DCs are impaired.28 Tregs limit the CD80/CD86–CD28-dependent activation of T cells through CTLA-4-dependent downregulation of CD80 and CD86 expression on DCs,29 which may contribute to the reduction in atherosclerosis.30 Therefore, we examined the interactions between CCR4-intact or CCR4-deficient Tregs and DCs by a coculture experiment (Figure 5B). As expected, the upregulation of CD80 and CD86 expression on DCs mediated by conventional T cells was markedly suppressed by coculture with Tregs from Apoe-/- or Ccr4-/-Apoe-/- mice (Figure 5C). Notably, the suppressive effect of CCR4-deficient Tregs was significantly attenuated compared with that of CCR4-intact Tregs (Figure 5C). To assess the involvement of CCR4 expression on Tregs in regulating the production of proatherogenic IFN-γ by conventional T cells, we stimulated conventional T cells with an anti-CD3 antibody and DCs in the presence or absence of CCR4-intact or CCR4-deficient Tregs and analyzed IFN-γ production by ELISA (Figure 5B). IFN-γ production by conventional T cells was markedly suppressed by coculture with CCR4-intact Tregs, while it was not significantly affected by coculture with CCR4-deficient Tregs (Figure 5D). Importantly, there was a marked difference in the ability to suppress the production of IFN-γ from conventional T cells between CCR4-intact Tregs and CCR4-deficient Tregs (Figure 5D). These results suggest that the augmented Th1 cell responses in the peripheral lymphoid tissues of Ccr4-/-Apoe-/- mice are partly due to impaired Treg-dependent regulation of DC function.

Earlier version: As described above, CD4+Foxp3+ Treg accumulation in the atherosclerotic aorta of Ccr4−/−Apoe−/− mice was not altered despite the marked expansion of CD4+Foxp3+ Tregs in peripheral lymphoid tissues. Based on these findings and previous reports showing the impaired migratory capacity of CCR4-deficient Tregs to inflammatory tissues,17,18 we hypothesized that CCR4-deficient Tregs might have less capacity to migrate to atherosclerotic lesions. To address this issue, we performed a Treg transfer experiment using Apoe−/− or Ccr4−/−Apoe−/− mice on a Kaede-Tg background31 which provided us with the opportunity to faithfully track transferred Tregs by monitoring the fluorescent Kaede protein (Figure 5E). There was no difference in the migratory capacity of CCR4-intact or CCR4-deficient Kaede-expressing Tregs to the peripheral LNs, spleen, or para-aortic LNs of recipient Apoe−/− mice (Figure 5F through 5H). To promote the accumulation of T cells in atherosclerotic aorta, we fed recipient Apoe−/− mice with a high-cholesterol diet and analyzed the migration of transferred Tregs in the aorta.

Revised version: As described above, CD4+Foxp3+ Treg accumulation in the atherosclerotic aorta of Ccr4-/-Apoe-/- mice was not altered despite the marked expansion of CD4+Foxp3+ Tregs in peripheral lymphoid tissues. Based on these findings and previous reports showing the impaired migratory capacity of CCR4-deficient Tregs to inflammatory tissues,17,18 we hypothesized that CCR4-deficient Tregs might have less capacity to migrate to atherosclerotic lesions. To address this issue, we performed a Treg transfer experiment using Apoe-/- or Ccr4-/-Apoe-/- mice on a Kaede-Tg background31 which provided us with the opportunity to faithfully track transferred Tregs by monitoring the fluorescent Kaede protein (Figure 5E). There was no difference in the migratory capacity of CCR4-intact or CCR4-deficient Kaede-expressing Tregs to the peripheral LNs, spleen, or para-aortic LNs of recipient Apoe-/- mice (Figure 5F through 5H). To promote the accumulation of T cells in atherosclerotic aorta, we fed recipient Apoe-/- mice with a high-cholesterol diet and analyzed the migration of transferred Tregs in the aorta. Interestingly, we found a trend toward reduction in the proportion of CCR4-deficient Kaede-expressing Tregs in the aorta of recipient Apoe-/- mice (Figure 5I), suggesting that CCR4-deficient Tregs may have a reduced ability to migrate to the atherosclerotic aorta, but not to the peripheral lymphoid tissues, under hypercholesterolemia.

Removed in the revised version: Interestingly, we found a significant reduction in the proportion of CCR4-deficient Kaede-expressing Tregs in the aorta of recipient Apoe−/− mice (Figure 5I), suggesting that CCR4-deficient Tregs have a reduced ability to migrate to the atherosclerotic aorta, but not to the peripheral lymphoid tissues, under hypercholesterolemia.

Earlier version: Taken together, these data demonstrate that CCR4 expression on Tregs plays a critical role in regulating Th1 cell responses in lymphoid tissues and mediating Treg migration to the atherosclerotic aorta under hypercholesterolemia, which cooperatively contributes to the reduction in early atherosclerosis by efficiently mitigating aortic inflammatory immune responses.

Revised version: Taken together, these data demonstrate that CCR4 expression on Tregs plays a critical role in regulating Th1 cell responses in lymphoid tissues and may mediate Treg migration to the atherosclerotic aorta under hypercholesterolemia, which may cooperatively contribute to the reduction in early atherosclerosis by efficiently mitigating aortic inflammatory immune responses.

CCR4 expression on Tregs is critical for limiting aortic inflammation and the development of atherosclerosis

Earlier version: To provide direct evidence for the critical role of CCR4 expression in Tregs in reducing early atherosclerosis, we injected Apoe−/− mice with saline or Tregs from Apoe−/− or Ccr4−/−Apoe−/− mice and analyzed the aortic root atherosclerotic lesions of recipient Apoe−/− mice (Figure 6A). Although there were no significant differences in the aortic sinus mean plaque area among the 3 groups (data not shown), detailed analysis of the aortic root plaques at 5 different levels revealed significantly greater lesions in Apoe−/− mice that were administered CCR4-deficient Tregs than in Apoe−/− mice that were administered CCR4-intact Tregs (Figure 6B), suggesting that the impaired CCR4-deficient Treg function is involved in the acceleration of atherosclerosis in Ccr4−/−Apoe−/− mice. There were no significant differences in body weight or plasma lipid profile among the 3 groups (Supplementary Table 2). We further aimed to confirm these findings by performing an additional experiment in which recipient Ccr4−/−Apoe−/− mice were administered Tregs from Apoe−/− or Ccr4−/−Apoe−/− mice or PBS and atherosclerotic lesions were analyzed (Supplementary Figure 8A). There were no major differences in body weight or plasma lipid profile among the 3 groups (Supplementary Table 3). The anti-atherogenic effect of CCR4 expression in Tregs was not observed in these mice (Supplementary Figure 8B). This could be possibly explained by the dysfunction of Tregs under the enhanced inflammatory conditions in Ccr4−/−Apoe−/− mice.

Revised version: To provide direct evidence for the critical role of CCR4 expression in Tregs in reducing early atherosclerosis, we injected Apoe-/- mice with saline or Tregs from Apoe-/- or Ccr4-/-Apoe-/- mice and analyzed the aortic root atherosclerotic lesions of recipient Apoe-/- mice (Figure 6A). Although there were no significant differences in the aortic sinus mean plaque area among the 3 groups (data not shown), detailed analysis of the aortic root plaques at 5 different levels revealed significantly greater lesions in Apoe-/- mice that were administered CCR4-deficient Tregs than in Apoe-/- mice that were administered CCR4-intact Tregs (Figure 6B), suggesting that the impaired CCR4-deficient Treg function is involved in the acceleration of atherosclerosis in Ccr4-/-Apoe-/- mice. There were no significant differences in body weight or plasma lipid profile among the 3 groups (Supplementary Table 2). We further aimed to confirm these findings by performing an additional experiment in which recipient Ccr4-/-Apoe-/- mice were administered Tregs from Apoe-/- or Ccr4-/-Apoe-/- mice or PBS and atherosclerotic lesions were analyzed (Supplementary Figure 10A). There were no major differences in body weight or plasma lipid profile among the 3 groups (Supplementary Table 3). The anti-atherogenic effect of CCR4 expression in Tregs was not observed in these mice (Supplementary Figure 10B). This could be possibly explained by the dysfunction of Tregs under the enhanced inflammatory conditions in Ccr4-/-Apoe-/- mice.

C-C chemokine receptor 4 (CCR4) expression on regulatory T cells (Tregs) is critical for limiting aortic inflammation and the development of early atherosclerosis.

A, Tregs purified from the peripheral lymph nodes (LNs) and spleen of apolipoprotein E-deficient (Apoe-/-) mice or CCR4-deficient mice on an Apoe-/- background (Ccr4-/-Apoe-/-) were intravenously transferred into 12-week-old Apoe-/- mice fed a standard chow diet, and atherosclerotic lesions were analyzed at 16 weeks of age. As a control without cell transfer, 12-week-old Apoe-/- mice were intravenously injected with saline and atherosclerotic lesions were analyzed at 16 weeks of age. B, Representative photomicrographs of Oil Red O staining and quantitative analysis of atherosclerotic lesion area at 5 different levels and maximal lesions in the aortic sinus of Apoe-/- mice injected with saline (n=19), Apoe-/- Tregs (n=17), or Ccr4-/-Apoe-/- Tregs (n=20). C-E, Representative sections and quantitative analyses of MOMA-2+ macrophages (C), CD4+ T cells (D), and collagen (E) in the aortic sinus of Apoe-/- mice injected with saline, Apoe-/- Tregs, or Ccr4-/-Apoe-/- Tregs. Arrowheads indicate the CD4+ T cells. n=10 per group. Black bars represent 50, 200, or 500 μm as described. Data points represent individual animals. Horizontal bars represent means. Error bars indicate s.d. *P<0.05, **P<0.01; 1-way ANOVA followed by Tukey’s multiple comparisons test.

Earlier version: To determine the effect of Treg-specific CCR4 deficiency on plaque components, we performed immunohistochemical studies of atherosclerotic lesions in the aortic sinus. Compared with those of saline-injected Apoe−/− mice, the atherosclerotic lesions of Apoe−/− mice injected with CCR4-intact Tregs showed markedly reduced accumulation of macrophages (Figure 6C) and CD4+ T cells (Figure 6D), whereas there were no differences in intraplaque accumulation of these inflammatory cells between saline-injected Apoe−/− mice and Apoe−/− mice injected with CCR4-deficient Tregs (Figure 6C and 6D). Notably, macrophage accumulation in the aortic sinus atherosclerotic lesions was markedly higher in Apoe−/− mice injected with CCR4-deficient Tregs than in Apoe−/− mice injected with CCR4-intact Tregs (Figure 6C). Collagen content in atherosclerotic lesions was significantly higher in Apoe−/− mice injected with CCR4-intact Tregs than in saline-injected Apoe−/− mice, while no difference in collagen was observed between between saline-injected Apoe−/− mice and Apoe−/− mice injected with CCR4-deficient Tregs (Figure 6E).

Revised version: To determine the effect of Treg-specific CCR4 deficiency on plaque components, we performed immunohistochemical studies of atherosclerotic lesions in the aortic sinus. Compared with those of saline-injected Apoe-/- mice, the atherosclerotic lesions of Apoe-/- mice injected with CCR4-intact Tregs showed markedly reduced accumulation of macrophages (Figure 6C) and CD4+ T cells (Figure 6D), whereas there were no differences in intraplaque accumulation of these inflammatory cells between saline-injected Apoe-/- mice and Apoe-/- mice injected with CCR4-deficient Tregs (Figure 6C and 6D). Notably, macrophage accumulation in the aortic sinus atherosclerotic lesions was markedly higher in Apoe-/- mice injected with CCR4-deficient Tregs than in Apoe-/- mice injected with CCR4-intact Tregs (Figure 6C). Collagen content in atherosclerotic lesions was significantly higher in Apoe-/- mice injected with CCR4-intact Tregs than in saline-injected Apoe-/- mice, while no difference in collagen was observed between saline-injected Apoe-/- mice and Apoe-/- mice injected with CCR4-deficient Tregs (Figure 6E).

Overall, we provide evidence that CCR4 protects against early atherosclerosis partly by mediating Treg-dependent induction of a less inflammatory plaque phenotype.

Discussion

Earlier version: In the present study, we demonstrated that genetic deletion of CCR4 in hypercholesterolemic Apoe−/− mice accelerates the development of early atherosclerotic lesions in the aortic root which exhibit an inflammatory phenotype, associated with the augmentation of proatherogenic Th1 cell responses in the peripheral lymphoid tissues, para-aortic LNs, and atherosclerotic aorta. Furthermore, T cell-DC coculture and Treg transfer experiments revealed that CCR4 expression on Tregs regulates the development of early atherosclerosis by suppressing Th1 cell responses in lymphoid tissues and mediating Treg migration to the atherosclerotic aorta. Thus, we identified a novel role for the CCL17/CCL22–CCR4 axes in controlling early atherosclerosis via favorable modulation of the Th1 cell/Treg balance.

Revised version: In the present study, we demonstrated that genetic deletion of CCR4 in hypercholesterolemic Apoe-/- mice accelerates the development of early atherosclerotic lesions in the aortic root which exhibit an inflammatory phenotype, associated with the augmentation of proatherogenic Th1 cell responses in the peripheral lymphoid tissues, para-aortic LNs, and atherosclerotic aorta. Furthermore, T cell-DC coculture and Treg transfer experiments revealed that CCR4 expression on Tregs regulates the development of early atherosclerosis by suppressing Th1 cell responses in lymphoid tissues and possibly by mediating Treg migration to the atherosclerotic aorta. Thus, we identified a novel role for the CCL17/CCL22–CCR4 axes in controlling early atherosclerosis via favorable modulation of the Th1 cell/Treg balance.

Emerging experimental and clinical data obtained by single-cell RNA sequencing and mass cytometry have clearly shown that CD4+ or CD8+ T cells are the dominant populations in human5 and mouse32 atherosclerotic plaques. Recent experimental and clinical evidence has demonstrated that dysregulation of the balance of proatherogenic Th1 cells and anti-atherogenic Tregs has adverse effects on atherosclerotic disease.13 In the present study, we found that the formation of early atherosclerotic lesions in the aortic root was markedly accelerated by CCR4 deficiency, which was associated with the upregulation of various helper T cell immune responses in peripheral lymphoid tissues and augmented Th1 and Th17 cell-mediated responses in the atherosclerotic aorta. Given the proatherogenic nature of Th1 cells,7,8 our data indicate that augmented Th1 cell responses in peripheral lymphoid tissues, especially in atherosclerotic lesions, substantially contributed to atherosclerotic lesion formation in Ccr4-/-Apoe-/- mice. However, we cannot exclude the possibility that augmented immune responses of other helper T cell subsets might also contribute to accelerated atherosclerosis in Ccr4-/-Apoe-/- mice.

Earlier version: As Th2 cells and some Th17 cells, but not Th1 cells, are known to express CCR4,15 the augmented Th1 cell responses observed in Ccr4−/−Apoe−/− mice do not seem to be direct effects of CCR4 deficiency and may involve other mechanisms. Given our recent report showing the augmented immune responses related to Th1, Th2, and Th17 cells in CD4+Foxp3+ Treg-depleted atherosclerosis-prone mice,33 augmentation of these helper T cell responses may be caused by dysregulated CD4+Foxp3+ Treg responses in our Ccr4−/−Apoe−/− mice. In line with experimental studies in mouse models of inflammatory diseases showing a pivotal role for CCR4 in mediating Treg migration to inflammatory tissues,17,18 we found that CCR4 expression on Tregs is critical for regulating atherosclerosis by mediating their migration to the atherosclerotic aorta. Another interesting finding is that CCR4 deficiency markedly upregulated Th1 cell responses in peripheral lymphoid tissues as well as in atherosclerotic lesions under hypercholesterolemia. T cell-DC coculture experiments clearly demonstrated that CCR4 expression on Tregs is critical for attenuating DC-dependent stimulation of Th1 cell responses by maintaining interactions between Tregs and DCs and Treg suppressive function, which is supported by the findings of a previous report showing a crucial role of the CCL22–CCR4 axis in the contact of Tregs with DCs and regulation of inflammatory responses.28 Based on these findings, we propose that CCR4 deficiency in Tregs interrupts their contact with DCs in lymphoid tissues and impairs their suppressive function despite their unaltered activation status, leading to augmented Th1 cell responses and accelerated atherosclerosis. We speculate that the increased frequency of Tregs in the peripheral lymphoid tissues of Ccr4−/−Apoe−/− mice may be a compensatory effect in response to augmented Th1 cell-mediated proinflammatory responses due to dysfunctional Tregs.

Revised version: As Th2 cells and some Th17 cells, but not Th1 cells, are known to express CCR4,15 the augmented Th1 cell responses observed in Ccr4-/-Apoe-/- mice do not seem to be direct effects of CCR4 deficiency and may involve other mechanisms. Given our recent report showing the augmented immune responses related to Th1, Th2, and Th17 cells in CD4+Foxp3+ Treg-depleted atherosclerosis-prone mice,33 augmentation of these helper T cell responses may be caused by dysregulated CD4+Foxp3+ Treg responses in our Ccr4-/-Apoe-/- mice. In line with experimental studies in mouse models of inflammatory diseases showing a pivotal role for CCR4 in mediating Treg migration to inflammatory tissues,17,18 we found that CCR4 expression on Tregs may be critical for their migration to the atherosclerotic aorta, providing a possible mechanism for CCR4-dependent regulation of atherosclerosis. Another interesting finding is that CCR4 deficiency markedly upregulated Th1 cell responses in peripheral lymphoid tissues as well as in atherosclerotic lesions under hypercholesterolemia. This observation may not depend on the activation status of Tregs because there were no major differences in the expression of their activation- and function-associated molecules between Apoe-/- and Ccr4-/-Apoe-/- mice. Despite the lower expression of CD103 in the peripheral LN Tregs of 8-week-old Ccr4-/-Apoe-/- mice, there was no significant difference in its expression levels between 18-week-old Apoe-/- and Ccr4-/-Apoe-/- mice, indicating that reduced CD103 expression in Ccr4-/-Apoe-/- mice may not be a noteworthy change. T cell-DC coculture experiments clearly demonstrated that CCR4 expression on Tregs is critical for attenuating DC-dependent stimulation of Th1 cell responses by maintaining interactions between Tregs and DCs and Treg suppressive function, which is supported by the findings of a previous report showing a crucial role of the CCL22–CCR4 axis in the contact of Tregs with DCs and regulation of inflammatory responses.28 In line with the results of in vitro experiments, we found the upregulation of CD86 expression on DCs in Ccr4-/-Apoe-/- mice, which may be derived from impaired ability of CCR4-deficient Tregs to downregulate CD80 and CD86 expression on DCs. Based on these findings, we propose that CCR4 deficiency in Tregs interrupts their contact with DCs in lymphoid tissues and impairs their suppressive function, leading to augmented Th1 cell responses and accelerated atherosclerosis. We speculate that the increased frequency of Tregs in the peripheral lymphoid tissues of Ccr4-/-Apoe-/- mice may be a compensatory effect in response to augmented Th1 cell-mediated proinflammatory responses due to dysfunctional Tregs.

Earlier version: T cell costimulatory signals such as the CD28–CD80/CD8610 and CD27–CD7034 pathways protect against atherosclerosis by systemically shifting the Th1 cell/Treg balance toward Treg responses. Several approaches involving treatment with cytokines35,36, antibodies11,36,37, an active form of vitamin D338, or ultraviolet B irradiation39,40 have been reported to mitigate atherosclerosis via modulation of the Th1 cell/Treg balance in the peripheral lymphoid tissues and atherosclerotic lesions of atherosclerosis-prone mice. Despite these reports, the mechanisms that regulate the T cell balance in lymphoid tissues, particularly in atherosclerotic lesions, have not been fully elucidated. The chemokine system plays an important role in the differentiation and migration of various subsets of T cells including Tregs and has been shown to be involved in the process of atherosclerosis.14 We revealed that CCR4 protects against early atherosclerosis by mediating Treg migration to the aorta and mitigating Th1 cell responses in lymphoid tissues and atherosclerotic lesions, indicating a previously unrecognized role of the chemokine system in regulating the Th1 cell/Treg balance to limit atherosclerosis.

Revised version: T cell costimulatory signals such as the CD28–CD80/CD8610 and CD27–CD7034 pathways protect against atherosclerosis by systemically shifting the Th1 cell/Treg balance toward Treg responses. Several approaches involving treatment with cytokines35,36, antibodies11,36,37, an active form of vitamin D338, or ultraviolet B irradiation39,40 have been reported to mitigate atherosclerosis via modulation of the Th1 cell/Treg balance in the peripheral lymphoid tissues and atherosclerotic lesions of atherosclerosis-prone mice. Despite these reports, the mechanisms that regulate the T cell balance in lymphoid tissues, particularly in atherosclerotic lesions, have not been fully elucidated. The chemokine system plays an important role in the differentiation and migration of various subsets of T cells including Tregs and has been shown to be involved in the process of atherosclerosis.14 We revealed that CCR4 protects against early atherosclerosis by mitigating Th1 cell responses in lymphoid tissues and atherosclerotic lesions and possibly by mediating Treg migration to the aorta. This indicates a previously unrecognized role of the chemokine system in regulating the Th1 cell/Treg balance to limit atherosclerosis.

We observed the accelerated development of early atherosclerotic lesions in Ccr4-/-Apoe-/- mice fed a standard chow diet. However, previous studies showed that neither hematopoietic nor systemic CCR4 deficiency affected the development of advanced atherosclerotic lesions or Treg frequency in severely hypercholesterolemic mice.22,24 These findings suggest that the role of CCR4 may vary depending on the stages of atherosclerosis. Given that Treg responses affected by CCR4 deficiency in mice with early atherosclerotic lesions may be different from those in mice with advanced atherosclerotic lesions, we speculate that the difference in Treg responses may be responsible for the inconsistent findings on the role of CCR4 deficiency in atherosclerosis.

Our data show that interactions between CCR4 and CCL17/CCL22 may promote Treg-skewed responses in lymphoid tissues and atherosclerotic lesions. However, studies in hypercholesterolemic mice demonstrated that genetic deficiency of CCL17 ameliorated atherosclerotic lesion development by promoting Treg accumulation in lymphoid tissues and atherosclerotic lesions independently of the CCL17–CCR4 axis.22,24 This suggests that CCL17 restrains Treg homeostasis and accelerates atherosclerosis, although the role of the CCL17–CCR4 axis in the regulation of Treg homeostasis and atherosclerosis remains elusive. A study in mouse models of myocardial injury showed that CCL17 deficiency increased Treg recruitment to the heart and attenuated myocardial inflammation and injury, which may be attributed to the cancelation of competitive inhibition of CCL22-mediated Treg chemotaxis by CCL17,41 suggesting that CCL17 may impair Treg homeostasis via the inhibition of the CCL22–CCR4 axis. Given the crucial role of the CCL22–CCR4 axis in meditating the Treg suppressive function described above,28 we propose that the CCL22–CCR4 axis plays a dominant role in the prevention of atherosclerosis by promoting Treg function, which may be inhibited by stimulating the CCL17–CCR4 axis. However, the specific roles of the CCL17–CCR4 and CCL22–CCR4 axes in the various stages of atherosclerosis have not been fully elucidated and further studies are needed.

Earlier version: This study has several limitations. Although Treg transfer experiments revealed a critical role for CCR4 in Tregs in protecting against early atherosclerosis, the use of conditional knockout mice would provide additional definitive evidence. Our data were obtained from animal experiments, and investigations using human samples will be needed to translate our findings to clinical settings.

Revised version: This study has several limitations. In flow cytometric analysis, we did not exclude dead cells or doublets. This procedure could have increased the reliability of our data. Although Treg transfer experiments revealed a critical role for CCR4 in Tregs in protecting against early atherosclerosis, the use of conditional knockout mice would provide additional definitive evidence. Our data were obtained from animal experiments, and investigations using human samples will be needed to translate our findings to clinical settings.

Earlier version: In conclusion, we demonstrated that CCR4 protects against early atherosclerosis by favorably modulating the balance between proatherogenic Th1 cell responses and atheroprotective Treg responses. We showed that CCR4 expression on Tregs is critical for suppressing Th1 cell responses and mediating Treg migration to the atherosclerotic aorta. Our data suggest that CCR4 is an important negative regulator of atherosclerosis.

Revised version: In conclusion, we demonstrated that CCR4 protects against early atherosclerosis by favorably modulating the balance between proatherogenic Th1 cell responses and atheroprotective Treg responses. We showed that CCR4 expression on Tregs is critical for suppressing Th1 cell responses and may play an important role in mediating Treg migration to the atherosclerotic aorta. Our data suggest that CCR4 is an important negative regulator of atherosclerosis.

Methods

Animals

All mice were male on a C57BL/6 background and fed a standard chow diet or a high-cholesterol diet containing 1.25% cholesterol (CLEA Japan, Tokyo, Japan) as indicated. Wild-type mice were obtained from CLEA Japan. Apoe-/-11 and Ccr4-/- mice42 are previously described. We crossed Ccr4-/- mice with Apoe-/- mice to obtain Ccr4-/-Apoe-/- mice. Kaede-Tg mice (RBRC05737)31 were provided by the RIKEN BRC through the National BioResource Project of the MEXT/AMED, Japan. We housed mice in cages for each strain or treatment group in a specific pathogen-free animal facility at Kobe Pharmaceutical University. Randomization and allocation concealment were performed. Littermate mice of each genotype were randomly allocated to each experimental group. During the experiments, animal/cage location was not controlled. The investigators were not blinded to the mouse genotype or treatment allocation. The criterion for exclusion was defined as severe body weight loss and set before the study. However, during at least 2 observations per week, we did not find such symptoms, and no mice were excluded. The experimental procedures were performed in our laboratory rooms or animal facility. All animal experiments were approved by the Animal Care Committee of Kobe Pharmaceutical University (permit numbers: 2018-008, 2019-011, 2020-050, 2021-038, 2022-005, 2023-038, and 2024-014) and conformed to the National Institutes of Health Guide for the Care and Use of Laboratory Animals and the ARRIVE guidelines (Animal Research: Reporting of In Vivo Experiments).

Assessment of biochemical parameters

Under anesthesia by intraperitoneal injection of medetomidine hydrochloride (0.3 mg/kg), midazolam (4 mg/kg), and butorphanol tartrate (5 mg/kg) (all from WAKO, Osaka, Japan), blood was collected by the cardiac puncture after overnight fasting, and plasma lipid profile was analyzed as described previously.40 Concentrations of plasma total cholesterol, high-density lipoprotein-cholesterol, and triglycerides were determined enzymatically using an automated chemistry analyzer (Oriental Yeast Co., Ltd., Tokyo, Japan).

Assessment of atherosclerotic lesion

The atherosclerotic lesions in the aortic root and thoracoabdominal aorta were analyzed as described previously.40 For analysis of atherosclerotic lesions in the aortic root, 5 consecutive sections (10 μm thickness), spanning 600 μm of the aortic sinus (150 μm interval), were collected from each mouse and stained with hematoxylin-eosin. The lesion area of the sections was quantified using the ImageJ (National Institutes of Health). Some sections were stained with Oil Red O (Sigma) for representative photomicrographs of the aortic sinus atherosclerotic lesions. For en face analysis of thoracoabdominal aortas, the aorta was opened longitudinally and stained with Oil Red O. The proportion of the lesion area was determined using the ImageJ.

Histological analysis of atherosclerotic lesions

Immunofluorescence staining of CCL17 and CCL22 in lymphoid tissues and atherosclerotic lesions was performed on 4% paraformaldehyde-fixed cryosections of mouse aortic roots using rabbit anti-CCL17 (1:200; abcam) or goat anti-CCL22 (1:200; R&D Systems) antibodies, followed by detection with fluorescent secondary antibodies. For costaining of macrophages in the above cryosections, anti-MOMA-2 (1:400; BMA Biomedicals) and fluorescent secondary antibodies were also used. Stained sections were digitally captured using a fluorescence microscope (BZ-X810; KEYENCE, Osaka, Japan).

For the detection of macrophages or CD4+ T cells, immunohistochemistry was performed on 4% paraformaldehyde-fixed cryosections of mouse aortic roots using anti-MOMA-2 or anti-CD4 (1:100; BD Biosciences) antibodies, followed by detection with biotinylated secondary antibodies and streptavidin-horseradish peroxidase. Staining with Masson’s trichrome (Muto Pure Chemicals, Tokyo, Japan) was used to delineate the fibrous area. Stained sections were digitally captured using a microscope (BZ-X810; KEYENCE), and the percentage of the stained area (the stained area per total atherosclerotic lesion area) was calculated as described previously.11 CD4+ T cells were quantified as described previously by counting the number of positively stained cells, which was divided by the total plaque area.11 The primary and secondary antibodies used are listed in Supplementary Table 4.

Flow cytometry

Earlier version: For flow cytometric analysis of lymphoid tissues, peripheral LN cells and splenocytes were isolated and stained in PBS containing 2% fetal calf serum. For analysis of immune cells within the aorta, mice were anesthetized, and the aorta was perfused with cold saline. The aorta was dissected and the adventitial tissue was carefully removed. The aorta was digested with Liberase TM (Roche Diagnostics) in plane RPMI medium at 37[ for 45 min with vortexing. For the detection of CCR4 on aortic T cells, Collagenase D (Sigma-Aldrich) was used instead of Liberase TM. A cell suspension obtained by mashing the aorta through a 70-μm strainer was stained with antibodies specific for CD3, CD4, CD45, CCR4, Foxp3, T-bet, RORγt, and GATA3.

Revised version: For flow cytometric analysis of lymphoid tissues, peripheral LN cells and splenocytes were isolated and stained in PBS containing 2% fetal calf serum. For analysis of immune cells within the aorta, mice were anesthetized, and the aorta was perfused with cold saline. The aorta was dissected and the adventitial tissue was carefully removed. The aorta was digested with Liberase TM (Roche Diagnostics) in plane RPMI medium at 37℃ for 45 min with vortexing. For the detection of CCR4 on aortic T cells, Collagenase D (Sigma-Aldrich) was used instead of Liberase TM. A cell suspension obtained by mashing the aorta through a 70-μm strainer was stained with antibodies specific for CD3, CD4, CD45, CCR4, Foxp3, T-bet, RORγt, and GATA3. The Foxp3 staining buffer set (Thermo Fisher Scientific) was used for intracellular staining of Foxp3. In all staining procedures, Fc receptors were blocked by anti-CD16/CD32 (BD Biosciences). Flow cytometric analysis was performed with FACSAria III (BD Biosciences) using FlowJo software version 10.8.1 (Tree Star). The antibodies used were listed in Supplementary Table 5. Gating strategy of flow cytometric analysis of aortic T cells was shown in Supplementary Figure 11.

Removed in the revised version: The Foxp3 staining buffer set (Thermo Fisher Scientific) was used for intracellular staining of Foxp3. In all staining procedures, Fc receptors were blocked by anti-CD16/CD32 (BD Biosciences). Flow cytometric analysis was performed with FACSAria III (BD Biosciences) using FlowJo software version 10.8.1 (Tree Star). The antibodies used were listed in Supplementary Table 5. Gating strategy of flow cytometric analysis of aortic T cells was shown in Supplementary Figure 9.

Intracellular cytokine staining

Immune cells from lymphoid tissues were stimulated with 20 ng/ml phorbol 12-myristate 13-acetate (Sigma) and 1 mmol/L ionomycin (Sigma) for 5 hours in the presence of Brefeldin A (Thermo Fisher Scientific). Intracellular cytokine staining was performed as described previously.40

Cytokine assay

In cell culture experiments, RPMI 1640 medium (Sigma) supplemented with 10% fetal calf serum, 50 μmol/L 2β-mercaptoethanol, and antibiotics was used. The production of several major cytokines from CD4+ T cells was examined as described previously.30 Splenic CD4+ T cells (1×105 cells) isolated using MACS (Miltenyi Biotec) were stimulated with plate-bound anti-CD3 (10 µg/mL, clone 145-2C11; BD Biosciences) and soluble anti-CD28 antibodies (2 µg/mL, clone 37.51; BD Biosciences) in 96-well round-bottomed plates for 48 hours. The concentrations of IL-4, IL-10, IL-17, and IFN-γ in culture supernatants were determined by ELISA using paired antibodies specific for corresponding cytokines (R&D Systems). The levels of multiple inflammation-associated cytokines and chemokines in culture supernatants were determined using a Mouse Cytokine Array Kit according to the manufacturer’s instructions (R&D Systems).

Earlier version: In some experiments, CD4+CD25+ Tregs and CD4+CD25− T cells were purified from the peripheral LNs and spleen of Apoe−/− or Ccr4−/−Apoe−/− mice using a CD4+CD25+ Regulatory T Cell Isolation Kit (Miltenyi Biotec) and anti-CD4 beads (Miltenyi Biotec) according to the manufacturer’s instructions. The purity of each population was >95% and most of the CD4+CD25+ T cells expressed Foxp3, as determined by flow cytometric analysis. The viability of each cell population was analyzed using 7-amino-actinomycin D (BD Biosciences) and Alexa Fluor 488 Annexin V/Dead cell Apoptosis Kit (Thermo Fisher scientific). The detailed data were shown in Supplementary Figure 10. CD11c+ DCs were isolated from spleen of Apoe−/− mice treated with Liberase TM using MACS (Miltenyi Biotec). The purity of the CD11c+ population was approximately 93%, as determined by flow cytometric analysis. CD4+CD25− T cells (1×105 cells) from Apoe−/− mice and splenic CD11c+ DC (2×104 cells) were cocultured with or without CD4+CD25+ Tregs (1.25×104 cells) from Apoe−/− or Ccr4−/−Apoe−/− mice in the presence of soluble anti-CD3 antibody (0.5 μg/mL) in 96-well round-bottomed plates. Culture supernatants were collected at 48 hours and analyzed by ELISA for IFN-γ as described above.

Revised version: In some experiments, CD4+CD25+ Tregs and CD4+CD25- T cells were purified from the peripheral LNs and spleen of Apoe-/- or Ccr4-/-Apoe-/- mice using a CD4+CD25+ Regulatory T Cell Isolation Kit (Miltenyi Biotec) and anti-CD4 beads (Miltenyi Biotec) according to the manufacturer’s instructions. The purity of each population was >95% and most of the CD4+CD25+ T cells expressed Foxp3, as determined by flow cytometric analysis. The viability of each cell population was analyzed using 7-amino-actinomycin D (BD Biosciences) and Alexa Fluor 488 Annexin V/Dead cell Apoptosis Kit (Thermo Fisher scientific). The detailed data were shown in Supplementary Figure 12. CD11c+ DCs were isolated from spleen of Apoe-/- mice treated with Liberase TM using MACS (Miltenyi Biotec). The purity of the CD11c+ population was approximately 93%, as determined by flow cytometric analysis. CD4+CD25- T cells (1×105 cells) from Apoe-/- mice and splenic CD11c+ DC (2×104 cells) were cocultured with or without CD4+CD25+ Tregs (1.25×104 cells) from Apoe-/- or Ccr4-/-Apoe-/- mice in the presence of soluble anti-CD3 antibody (0.5 μg/mL) in 96-well round-bottomed plates. Culture supernatants were collected at 48 hours and analyzed by ELISA for IFN-γ as described above.

Treg suppression assay

For analysis of the in vitro suppressive function of Tregs, CD4+CD25+ Tregs and CD4+CD25- T cells were purified from pooled peripheral LNs and spleen of Apoe-/- or Ccr4-/-Apoe-/- mice as described above. Purified CD4+CD25+ Tregs from Apoe-/- or Ccr4-/-Apoe-/- mice were cocultured with carboxyfluorescein diacetate succinimidyl ester (Thermo Fisher Scientific)-labeled CD4+CD25- conventional T cells (2.5×104 cells) from Apoe-/- mice at the indicated ratios in the presence of mitomycin C (WAKO)-treated antigen-presenting cells (5×104 cells) and soluble anti-CD3 antibody (0.5 μg/mL) in 96-well round-bottomed plates. Splenocytes were used as antigen-presenting cells. The cocultured cells were maintained at 37°C with 5% CO2 for 3 days. The proliferation of carboxyfluorescein diacetate succinimidyl ester-labeled CD4+CD25- conventional T cells was analyzed by flow cytometry.

Flow cytometric analysis of DC phenotypic changes

CD4+CD25+ Tregs and CD4+CD25- T cells were purified from peripheral LNs and spleen of Apoe-/- or Ccr4-/-Apoe-/- mice, and CD11c+ DCs were isolated from the spleen of Apoe-/- mice as described above. CD4+CD25- T cells (5×104 cell) from Apoe-/- mice or a mixture of CD4+CD25+ Tregs (5×104 cells) from Apoe-/- or Ccr4-/-Apoe-/- mice and CD4+CD25- T cells (5×104 cells) from Apoe-/- mice were cocultured with splenic CD11c+ DCs (4×104 cells) in the presence of soluble anti-CD3 antibody (0.1 μg/mL) in 96-well round-bottomed plates. After 48 hours of coculture, the cells were collected, stained with antibodies specific for IAb, CD11c, CD80, CD86, and 7-amino-actinomycin D, and analyzed using FACSAria III (BD Biosciences).

Quantitative reverse transcription PCR analysis

Using TRIzol reagent (Thermo Fisher Scientific), we extracted total RNA from aorta which was perfused with cold saline and subsequently soaked in RNA later (Thermo Fisher Scientific). After the isolation of CD4+CD25+ Tregs and CD4+CD25- T cells as described above, we extracted total RNA from the cells using an RNeasy Mini Kit (Qiagen). A PrimeScript RT reagent Kit (Takara, Shiga, Japan) was used for reverse transcription. Quantitative PCR analysis was conducted using a TB Green Ex Taq (Takara) and a StepOnePlus Real-Time PCR System (Thermo Fisher Scientific) according to the manufacturer’s instructions. The primers used are listed in Supplementary Table 6. Amplification reactions were performed in duplicate and fluorescence curves were analyzed with the included software. GAPDH was used as an endogenous control reference.

In vivo Treg homing assay

Earlier version: CD4+CD25+ Tregs were purified from the peripheral LNs and spleen of Apoe−/−Kaede-Tg or Ccr4−/−Apoe−/−Kaede-Tg mice as described above and were intravenously injected into recipient 18-week-old Apoe−/− mice (1×106 cells per mouse) on a high-cholesterol diet containing 1.25% cholesterol via the tail vein. At 20 hours after transfer, the Kaede+ Treg proportions in the peripheral LNs, spleen, para-aortic LNs, and aorta of Apoe−/− mice were analyzed by flow cytometry. Gating strategy of flow cytometric analysis of aortic T cells was shown in Supplementary Figure 11.

Revised version: CD4+CD25+ Tregs were purified from the peripheral LNs and spleen of Apoe-/-Kaede-Tg or Ccr4-/-Apoe-/-Kaede-Tg mice as described above and were intravenously injected into recipient 18-week-old Apoe-/- mice (1×106 cells per mouse) on a high-cholesterol diet containing 1.25% cholesterol via the tail vein. At 20 hours after transfer, the Kaede+ Treg proportions in the peripheral LNs, spleen, para-aortic LNs, and aorta of Apoe-/- mice were analyzed by flow cytometry. Gating strategy of flow cytometric analysis of aortic T cells was shown in Supplementary Figure 13.

Analysis of atherosclerotic lesions in Treg-transferred mice

To clarify the role of CCR4 expression in Tregs in regulating atherosclerosis, CD4+CD25+ Tregs were purified from the peripheral LNs and spleen of Apoe-/- or Ccr4-/-Apoe-/- mice as described above and were intravenously injected into 12-week-old recipient Apoe-/- or Ccr4-/-Apoe-/- mice (1×106 cells per mouse) on a standard chow diet via the tail vein. The atherosclerotic lesions were analyzed 4 weeks after transfer as described above.

Statistical analysis

Earlier version: Normality was assessed by Shapiro-Wilk normality test. Two-tailed Student’s t-test or Mann-Whitney U-test was used to detect significant differences between 2 groups when appropriate. One-way ANOVA followed by Tukey’s multiple comparisons test or 2-way ANOVA followed by Tukey’s multiple comparisons test was performed for multiple groups where appropriate. A value of P<0.05 was considered statistically significant. No data were excluded from the analysis. The investigators were not blinded to the data analysis. For statistical analysis, GraphPad Prism version 9.0 (GraphPad Software Inc.) was used.

Revised version: Normality was assessed by Shapiro-Wilk normality test. Two-tailed Student’s t-test, Mann-Whitney U-test, or one-sample t-test was used to detect significant differences between 2 groups when appropriate. One-way ANOVA followed by Tukey’s multiple comparisons test or 2-way ANOVA followed by Tukey’s multiple comparisons test was performed for multiple groups where appropriate. A value of P<0.05 was considered statistically significant. No data were excluded from the analysis. The investigators were not blinded to the data analysis. For statistical analysis, GraphPad Prism version 9.0 (GraphPad Software Inc.) was used.